---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Insert exercise title here

```yaml
type: NormalExercise
key: fec7f8f3e6
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}
import pandas as pd
pd.DataFrame([1, 2, 3])
```

`@sct`
```{python}
Ex().check_function('pandas.DataFrame').check_args('data').has_equal_value()
success_msg("Well done!")
```

---

## Exercise 1.2

```yaml
type: NormalExercise
key: f868f5d4e6
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
